Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **34 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Pedro"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - The Delta Median problem and Randomization

## Question 1 [time estimate: 16 minutes]

### Question 1a [time estimate: 7 minutes]

Write the worst sort function in the world, `worst_sort`. This function takes a list, and then randomly shuffles it until it happens to be in sorted order. Once it is in sorted order then the list is returned.


<!-- grade_id: cell-ead9c74054a1c5da, points: 0, grade: false, solution: true -->
```python
import random

def worst_sort(A):
    """
    Sort A in ascending order by randomly shuffling its 
    elements until they are in order.
    
    Parameters
    ----------
    A : list
        A list of numerical values
        
    Returns
    -------
    A : list
        Sorted list
        
    """
    while A != sorted(A):
        random.shuffle(A)
    
    return A
```

<!-- grade_id: cell-ff4db3a3e8b04515, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 1b [time estimate: 3 minutes]
What is the best case complexity of this algorithm?

<!-- grade_id: cell-afac80171b3b6232, points: 0, grade: true, solution: true -->
the best case complexity here is $\Theta(1)$. If we get an already sorted list, there's no need to do anything

### Question 1c [time estimate: 3 minutes]

What is the average case complexity?


<!-- grade_id: cell-9d52d34daa1e3478, points: 0, grade: true, solution: true -->
Since the probability of getting a sortted list when shuffling is $\frac{1}{n!}$ then the average case complexity would be $\Omega(n!)$

### Question 1d [time estimate: 3 minutes]

For what size lists is this a feasible method?

<!-- grade_id: cell-ab40f08768579798, points: 0, grade: true, solution: true -->
well, it will depend on the power of your machine and the other algorithm you're comparing. A factorial function goes ver high, very quickly. Comparing to a squared function, it will be higher, for any list bigger than 4 elements.

### [Optional] Question 1e [time estimate: 5 minutes]

Can you think of an even worse sorting method? In such a case, what would its complexity be? How big a list could you sort?

<!-- grade_id: cell-ac05806942341926, points: 0, grade: true, solution: true -->
I can think of a worst algorithm than this one

## Question 2 [time estimate: 18 minutes]

Approximate median finder. Given a list and $\tilde{\delta}$ (a number between 0 and 0.5), the approximate median finder function returns a number that is guaranteed to lie between the $(50-\tilde{\delta}/2)$% and $(50+\tilde{\delta}/2)$% percentiles, where $\tilde{\delta}=100\cdot \delta$. Implement such a function by randomly selecting an element from the list and testing whether or not it lies within the bounds. If it doesn’t then retry with a new random element, but only a limited number of retries are allowed (you decide the maximum number of retries.)

### Question 2a [time estimate: 5 minutes]
Complete the following function


<!-- grade_id: cell-2ab65512a6148d3c, points: 0, grade: false, solution: true -->
```python
import numpy as np

def check_approx_med(A, num, delta):
    """
    Given a list and a number in a list, determine whether it is 
    within (50-delta/2)% and(50+delta/2)% percentiles.
    
    X% percentile of a list is defined to be the smallest number in 
    the list that is as large as at least X% numbers in the list. 
    
    Parameters
    ----------
    A : list
        A list of numerical values
    num 
        The number of interest
    delta : float
        The error, lies between 0 and 0.5

    Returns
    -------
    isApproxMed : bool
        A boolean value indicating whether num is whithin the bound
    
    """
    a = np.array(A)
    
    d = 100*delta
    
    p = np.percentile(a, 50 - d/2)
    
    Bp = np.percentile(a, 50 + d/2)
    
    return num >= p and num <= Bp
```

<!-- grade_id: cell-df56d7433abb6517, points: 1, grade: true, solution: false -->
```python
assert(check_approx_med([0, 1], 0.4, .25) == True)
assert(check_approx_med([0, 1], 0, .25) == False)
assert(check_approx_med([0], 0, .25) == True)
```

### Question 2b [time estimate: 5 minutes]
Complete the following function that makes use of `check_approx_med` above.


<!-- grade_id: cell-e19685b545147dde, points: 0, grade: false, solution: true -->
```python
import math

def approx_med_finder(A, delta):
    """
    Given a list, find a number in the list that is between 
    50-delta/2% and 50+delta/2% percentiles
    
    Paramters
    ---------
    A : list
        A list of numerical values
    Delta : int
        The error, lies between 0 and 0.5
        
    Returns
    -------
    num
        The approximated median if it is found within 100//delta trials
        (this is one possible maximum number of retries. While you are 
        encouraged to play around with another limits, the submitted work 
        must use this number)
    None
        Finding failed, if nothing is found within 100//delta trials
   
   
    Note
    ----
    The 100//delta is chosen here as it is the expected number of trials for 
    the first successful finding to occur (using geometric distribution)
    
    """
    trials = 100
    
    num = math.floor(random.random()*len(A))
    
    for i in range(trials):
        if check_approx_med(A, A[num], delta):
            return A[num]
        else:
            num = math.floor(random.random()*len(A))
    
    return None

approx_med_finder([1,1,1,1,1,1,1,1], 0.3)
```

Output:
```
1
```

### Question 2c [time estimate: 3 minutes]

What is the probability of failure in each random trial? What is the probability of failure after all the allowed random trials? Does this scale with $\delta$ or $N$ (the number of elements in a list)?


<!-- grade_id: cell-efe7ab76f6a2546a, points: 0, grade: true, solution: true -->
the probability scales with $\delta$, since the longer the list, if $\delta$ does not change, the bigger the amount of numbers within the percentile, meaning that the probability won't change

### Question 2d [time estimate: 5 minutes]
Analyze the expected runtime of `approx_med_finder`. Note that because the function uses `check_approx_med`, you will most likely need to analyze the runtime of that function, too.

<!-- grade_id: cell-86b199798477fc2a, points: 0, grade: true, solution: true -->
this function will depend a lot on the runtime of the percentile function in numpy, it will run a min of 2 times the runtime of percentile and a max of 2 times trials times the runtime of percentile

```python

```